---
title: 'Chapter 3'
description: 'Who are my most valuable customers?'
---

## Insert exercise title here

```yaml
type: NormalExercise
key: c6e27a9fa9
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```
